# Python review: Values, variables, types, lists, and strings

These first few notebooks are a set of exercises designed to reinforce various aspects of Python programming.

```python
### Global imports
import dill
from cse6040_devkit import plugins, utils
from cse6040_devkit.training_wheels import run_with_timeout, suppress_stdout
import tracemalloc
from time import time
import re 
import pandas as pd


```

### Exercise 0: (1 points)
**return_string__FREE**  

**Example:** we have defined `return_string__FREE` as follows:

**This is a free exercise!** 

**Please run the test cell below to collect your FREE point!**

Write a function that returns a string `Hello, world!`.


```python
### Solution - Exercise 0  
def return_string__FREE(string_to_return: str='Hello, world!') -> str:
    return string_to_return

### Demo function call
print(f'result={return_string__FREE()}')
```

Output:
```
result=Hello, world!
```

 

**The demo should display this printed output.**
```
result=Hello, world!
```


 ---
 <!-- Test Cell Boilerplate -->  
 The test cell below will always pass. Please submit to collect your free points for return_string__FREE (exercise 0).
 

<!-- grade_id: ex_0, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 0  


print('Passed! Please submit.')
```

Output:
```
Passed! Please submit.
```

### Exercise 1: (1 points)
**x_float_ex**  

**Your task:** define `x_float_ex` as follows:


Write a function that converts the integer `val` to a `float` and returns the corresponding value.

**Inputs**:
- `val` (int): an integer

**Return**:
- `x_float` (float): the input `val` represented as a float type


```python
### Solution - Exercise 1  
def x_float_ex(val: int=0) -> float:
    ###
    ### YOUR CODE HERE
    ###
    return float(val)

### Demo function call
print(f'result={x_float_ex(1)}')
```

Output:
```
result=1.0
```

 

**The demo should display this printed output.**
```
result=1.0
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for x_float_ex (exercise 1). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_1, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 1  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=x_float_ex,
              ex_name='x_float_ex',
              key=b'OiQUN6z33z2RTVEffvMxO0UCwxgRlYvOUyht3aXxIQg=', 
              n_iter=20)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to x_float_ex did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-7-19d21e64a459> in <module>
      6 from time import time
      7 
----> 8 tracemalloc.start()
      9 mem_start, peak_start = tracemalloc.get_traced_memory()
     10 print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

NameError: name 'tracemalloc' is not defined
```

### Exercise 2: (1 points)
**strcat_ba_test**  

**Your task:** define `strcat_ba_test` as follows:


Write a function that concatenates input strings `b` and then `a` and returns the resulting string.

**Inputs**:
- `a` (str)
- `b` (str)

**Return**:
- `b_a` (str): the concatenation of `b` and `a`


```python
### Solution - Exercise 2  
def strcat_ba_test(a: str, b: str) -> str:
    ###
    ### YOUR CODE HERE
    ###
    return b+a

### Demo function call
a='hello'
b='world'
print(f'result={strcat_ba_test(a,b)}')
```

Output:
```
result=worldhello
```

 

**The demo should display this printed output.**
```
result=worldhello
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for strcat_ba_test (exercise 2). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_2, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 2  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=strcat_ba_test,
              ex_name='strcat_ba_test',
              key=b'OiQUN6z33z2RTVEffvMxO0UCwxgRlYvOUyht3aXxIQg=', 
              n_iter=20)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to strcat_ba_test did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

### Exercise 3: (2 points)
**strcat_list_test**  

**Your task:** define `strcat_list_test` as follows:


Write a function that concatenates the input list of strings in reverse order and returns the resulting string.

**Inputs**:
- `L` (list)

**Return**:
- `b_a` (str): the concatenation of the list of strings in reverse order


```python
### Solution - Exercise 3  
def strcat_list_test(L: list) -> str:
    ###
    ### YOUR CODE HERE
    ###
    b_a = "".join(L[::-1])
    return b_a

### Demo function call
L=['!','world',', ','Hello']
print(f'result={strcat_list_test(L)}')
```

Output:
```
result=Hello, world!
```

 

**The demo should display this printed output.**
```
result=Hello, world!
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for strcat_list_test (exercise 3). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_3, points: 2, grade: true, solution: false -->
```python
### Test Cell - Exercise 3  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=strcat_list_test,
              ex_name='strcat_list_test',
              key=b'OiQUN6z33z2RTVEffvMxO0UCwxgRlYvOUyht3aXxIQg=', 
              n_iter=20)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to strcat_list_test did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-16-7f81e891b3fe> in <module>
      6 from time import time
      7 
----> 8 tracemalloc.start()
      9 mem_start, peak_start = tracemalloc.get_traced_memory()
     10 print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

NameError: name 'tracemalloc' is not defined
```

### Exercise 4: (1 points)
**floor_fraction_test**  

**Your task:** define `floor_fraction_test` as follows:


Write a function that returns $\left\lfloor\frac{a}{b}\right\rfloor$ also known as the *floor* of $\frac{a}{b}$.

**Inputs**:
- `a` (int): real non-negative number such that $a \geq 0$
- `b` (int): real positive number such that $b > 0$

**Return**:
- `floor_a_b` (int): the floor of $\frac{a}{b}$


```python
### Solution - Exercise 4  
def floor_fraction_test(a: int, b: int) -> int:
    def _is_number(x):
        from numbers import Number
        return isinstance(x,Number)
    assert _is_number(a) and a>=0
    assert _is_number(b) and b>0
    ###
    ### YOUR CODE HERE
    ###
    return int(a/b)

### Demo function call
a=7
b=3
print(f'result={floor_fraction_test(a,b)}')
```

Output:
```
result=2
```

 

**The demo should display this printed output.**
```
result=2
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for floor_fraction_test (exercise 4). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_4, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 4  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=floor_fraction_test,
              ex_name='floor_fraction_test',
              key=b'OiQUN6z33z2RTVEffvMxO0UCwxgRlYvOUyht3aXxIQg=', 
              n_iter=20)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to floor_fraction_test did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

### Exercise 5: (1 points)
**ceiling_fraction_test**  

**Your task:** define `ceiling_fraction_test` as follows:


Write a function that returns $\left\lceil\frac{a}{b}\right\rceil$ also known as the *ceiling* of $\frac{a}{b}$.

**Inputs**:
- `a` (int): real non-negative number such that $a \geq 0$
- `b` (int): real positive number such that $b > 0$

**Return**:
- `ceiling_a_b` (int): the ceiling of $\frac{a}{b}$


```python
### Solution - Exercise 5  
def ceiling_fraction_test(a: int, b: int) -> int:
    def _is_number(x):
        from numbers import Number
        return isinstance(x,Number)
    assert _is_number(a) and a>=0
    assert _is_number(b) and b>0
    ###
    ### YOUR CODE HERE
    ###
    return int((a + b -1)/b)

### Demo function call
a=7
b=3
print(f'result={ceiling_fraction_test(a,b)}')
```

Output:
```
result=3
```

 

**The demo should display this printed output.**
```
result=3
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for ceiling_fraction_test (exercise 5). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_5, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 5  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=ceiling_fraction_test,
              ex_name='ceiling_fraction_test',
              key=b'OiQUN6z33z2RTVEffvMxO0UCwxgRlYvOUyht3aXxIQg=', 
              n_iter=20)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to ceiling_fraction_test did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

### Exercise 6: (1 points)
**report_exam_avg_test**  

**Your task:** define `report_exam_avg_test` as follows:


Write a function that returns the average of score of three exam scores represented by `a`, `b`, and `c`.

**Inputs**:
- `a` (int): Exam Score `a`
- `b` (int): Exam Score `b`
- `c` (int): Exam Score `c`

**Return**:
- `avg_exam_str` (str): `Your average score: XX`, where `XX` is the average of 3 exams rounded to one decimal place


```python
### Solution - Exercise 6  
def report_exam_avg_test(a: int, b: int, c: int) -> str:
    ###
    ### YOUR CODE HERE
    ###
    avg_exam_str = f"Your average score: {round(sum([a, b, c]) / 3,1)}"

    return avg_exam_str
    

### Demo function call
a=100
b=95
c=80
print(f'result={report_exam_avg_test(a,b,c)}')
```

Output:
```
result=Your average score: 91.7
```

 

**The demo should display this printed output.**
```
result=Your average score: 91.7
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for report_exam_avg_test (exercise 6). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_6, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 6  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=report_exam_avg_test,
              ex_name='report_exam_avg_test',
              key=b'OiQUN6z33z2RTVEffvMxO0UCwxgRlYvOUyht3aXxIQg=', 
              n_iter=20)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to report_exam_avg_test did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

### Exercise 7: (2 points)
**count_word_lengths_test**  

**Your task:** define `count_word_lengths_test` as follows:


Write a function that returns a list containing the lengths of words in a string

**Inputs**:
- `s` (str): a phrase of words (lowercase alphabetic characters) separated by multiple character spaces

**Return**:
- `word_lengths` (list): Contains the lengths of words from input. If input is empty, function should return an empty list.


```python
### Solution - Exercise 7  
def count_word_lengths_test(s: str) -> list:
    ###
    ### YOUR CODE HERE
    ###
    word_lengths=[len(word) for word in s.split()]
    return word_lengths

### Demo function call
s='the quick  brown   fox jumped over     the lazy  dog'
print(f'result={count_word_lengths_test(s)}')
```

Output:
```
result=[3, 5, 5, 3, 6, 4, 3, 4, 3]
```

 

**The demo should display this printed output.**
```
result=[3, 5, 5, 3, 6, 4, 3, 4, 3]
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for count_word_lengths_test (exercise 7). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_7, points: 2, grade: true, solution: false -->
```python
### Test Cell - Exercise 7  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=count_word_lengths_test,
              ex_name='count_word_lengths_test',
              key=b'OiQUN6z33z2RTVEffvMxO0UCwxgRlYvOUyht3aXxIQg=', 
              n_iter=50)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to count_word_lengths_test did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

### FIN!
You've reached the end of this part. Don't forget to restart and run all cells again to make sure it's all working when run in sequence; and make sure your work passes the submission process. Good luck!