# Python review: Basic collections of values
This notebook continues the review of Python basics. The focus here is on basic collections: tuples, dictionaries, and sets.

```python
### Global imports
import dill
from cse6040_devkit import plugins, utils
from cse6040_devkit.training_wheels import run_with_timeout, suppress_stdout
import tracemalloc
from time import time
import re 
import pandas as pd


```

### Exercise 0: (1 points)
**minmax_test**  

**Your task:** define `minmax_test` as follows:


Write a function that returns a tuple containing the minimum and maximum values from the the input list.

**Inputs**:
- `L` (list): a list of integers

**Return**:
- `minmax` (tuple): a tuple where the first element is the minimum value from `L` and the second element is the maximum value from `L`


```python
### Solution - Exercise 0  
def minmax_test(L: list) -> tuple:
    ###
    ### YOUR CODE HERE
    ###
    minmax = tuple([min(L), max(L)])
        
    return minmax


### Demo function call
L=[8, 7, 2, 5, 1]
print(f'result={minmax_test(L)}')
```

Output:
```
result=(1, 8)
```

 

**The demo should display this printed output.**
```
result=(1, 8)
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for minmax_test (exercise 0). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_0, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 0  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=minmax_test,
              ex_name='minmax_test',
              key=b'5CpU4KJ6upxgu5f9QF_LK4KGaK5zkxGTtab-32LDWMo=', 
              n_iter=50)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to minmax_test did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

Output:
```
initial memory usage: 0.00 MB
Test duration: 0.17 seconds
memory after test: 3.10 MB
memory peak during test: 4.11 MB
Passed! Please submit.
```

### Exercise 1: (2 points)
**remove_all_test**  

**Your task:** define `remove_all_test` as follows:


Write a function that returns a *copy* of the input list `L` with all occurrences of `x` removed.

**Inputs**:
- `L` (list): a list of integers
- `x` (int): an integer that needs to be removed from `L`

**Return**:
- `removed_list` (list): a copied list of `L` where all occurences of `x` are removed.


```python
### Solution - Exercise 1  
def remove_all_test(L: list,x: int) -> list:
    ###
    ### YOUR CODE HERE
    ###
    removed_list = []
    for i in L:
        if i != x:
            removed_list.append(i)
                        
    return removed_list
    

### Demo function call
L=[1, 2, 3, 2, 4, 8, 2]
x=2
print(f'result={remove_all_test(L,x)}')
```

Output:
```
result=[1, 3, 4, 8]
```

 

**The demo should display this printed output.**
```
result=[1, 3, 4, 8]
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for remove_all_test (exercise 1). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_1, points: 2, grade: true, solution: false -->
```python
### Test Cell - Exercise 1  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=remove_all_test,
              ex_name='remove_all_test',
              key=b'5CpU4KJ6upxgu5f9QF_LK4KGaK5zkxGTtab-32LDWMo=', 
              n_iter=50)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to remove_all_test did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

Output:
```
initial memory usage: 0.00 MB
Test duration: 0.17 seconds
memory after test: 3.11 MB
memory peak during test: 4.21 MB
Passed! Please submit.
```

### Exercise 2: (2 points)
**compress_vector_test**  

**Your task:** define `compress_vector_test` as follows:


Write a function that returns a dictionary of indicies and non-zero values from a corresponding input list.

**Inputs**:
- `L` (list): a list of floats

**Return**:
- `d` (dict): 
  - `d['inds']` (list): the values corresponding to the indicies of the non-zero values from `L`
  - `d['vals']` (list): the values corresponding to the non-zero values from `L`


```python
### Solution - Exercise 2  
def compress_vector_test(L: list) -> dict:
    ###
    ### YOUR CODE HERE
    ###
    d = {'inds': [], 'vals': []}
    for i, val in enumerate(L):
        if val !=0:
            d['inds'].append(i)
            d['vals'].append(val)
    return d 

### Demo function call
L=[0.0, 0.87, 0.0, 0.0, 0.0, 0.32, 0.46, 0.0, 0.0, 0.10, 0.0, 0.0]
print(f'result={compress_vector_test(L)}')
```

Output:
```
result={'inds': [1, 5, 6, 9], 'vals': [0.87, 0.32, 0.46, 0.1]}
```

 

**The demo should display this printed output.**
```
result={'inds': [1, 5, 6, 9], 'vals': [0.87, 0.32, 0.46, 0.1]}
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for compress_vector_test (exercise 2). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_2, points: 2, grade: true, solution: false -->
```python
### Test Cell - Exercise 2  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=compress_vector_test,
              ex_name='compress_vector_test',
              key=b'5CpU4KJ6upxgu5f9QF_LK4KGaK5zkxGTtab-32LDWMo=', 
              n_iter=50)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to compress_vector_test did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

Output:
```
initial memory usage: 0.00 MB
Test duration: 0.16 seconds
memory after test: 3.10 MB
memory peak during test: 4.18 MB
Passed! Please submit.
```

### Repeated Indices 
Consider the compressed vector data structure, `d`, in the preceding exercise, which stores a list of indices (`d['inds']`) and a list of values (`d['vals']`).

Suppose we allow duplicate indices, possibly with different values. For example:

```python
    d['inds'] == [0,   3,   7,   3,   3,   5, 1]
    d['vals'] == [1.0, 2.0, 3.0, 4.0, 5.0, 6.0, 7.0]
```

In this case, the index 3 appears three times. (Also note that the indices `d['ind']` need not appear in sorted order.)

Let's adopt the convention that when there are repeated indices, the "true" value there is the _sum_ of the individual values. In other words, the true vector corresponding to this example of `d` would be:

```python
    # ind:  0    1    2    3*    4    5    6    7
    x == [1.0, 7.0, 0.0, 11.0, 0.0, 6.0, 0.0, 3.0]
```

### Exercise 3: (2 points)
**decompress_vector_test**  

**Your task:** define `decompress_vector_test` as follows:


Write a function that returns a list representing the vector based on a dictionary of indicies and values.

**Inputs**:
- `d` (dict):
  - `d['inds']` (list): the values corresponding to the indicies of the non-zero values
  - `d['vals']` (list): the values corresponding to the non-zero values
- `n` (int): the length of the full vector

**Return**:
- `L` (list): a list of floats

**Requirements**: 
- If there are repeated indicies, the values found in `d['vals']` should be summed


```python
### Solution - Exercise 3  
def decompress_vector_test(d: dict, n: int=None) -> list:
    # Checks the input
    assert type(d) is dict and 'inds' in d and 'vals' in d, "Not a dictionary or missing keys"
    assert type(d['inds']) is list and type(d['vals']) is list, "Not a list"
    assert len(d['inds']) == len(d['vals']), "Length mismatch"
    
    # Determine length of the full vector
    i_max = max(d['inds']) if d['inds'] else -1
    if n is None:
        n = i_max+1
    else:
        assert n > i_max, "Bad value for full vector length"
    ###
    ### YOUR CODE HERE
    ###
    L = [0.0]*n
    for i, val in zip(d['inds'],d['vals']):
            L[i] += val
    return L
       
            
    

### Demo function call
d={'inds':[0,3,7,3,3,5,1],
   'vals':[1.0,2.0,3.0,4.0,5.0,6.0,7.0]
   }
print(f'result={decompress_vector_test(d)}')
```

Output:
```
result=[1.0, 7.0, 0.0, 11.0, 0.0, 6.0, 0.0, 3.0]
```

 

**The demo should display this printed output.**
```
result=[1.0, 7.0, 0.0, 11.0, 0.0, 6.0, 0.0, 3.0]
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for decompress_vector_test (exercise 3). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_3, points: 2, grade: true, solution: false -->
```python
### Test Cell - Exercise 3  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=decompress_vector_test,
              ex_name='decompress_vector_test',
              key=b'5CpU4KJ6upxgu5f9QF_LK4KGaK5zkxGTtab-32LDWMo=', 
              n_iter=50)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to decompress_vector_test did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

Output:
```
initial memory usage: 0.00 MB
Test duration: 0.11 seconds
memory after test: 0.04 MB
memory peak during test: 1.22 MB
Passed! Please submit.
```

### Exercise 4: (1 points)
**find_common_inds_test**  

**Your task:** define `find_common_inds_test` as follows:


Write a function that returns a sorted list of common indicies from two uncompressed vectors.

**Inputs**:
- `d1` (dict):
  - `d1['inds']` (list): the values corresponding to the indicies of the non-zero values
  - `d1['vals']` (list): the values corresponding to the non-zero values
- `d2` (dict):
  - `d2['inds']` (list): the values corresponding to the indicies of the non-zero values
  - `d2['vals']` (list): the values corresponding to the non-zero values

**Return**:
- `common_L` (list): a sorted list of common indicies

**Requirements**: 
- If there are repeated indicies, only return the index once


```python
### Solution - Exercise 4  
def find_common_inds_test(d1: dict, d2: dict) -> list:
    # Checks the input
    assert type(d1) is dict and 'inds' in d1 and 'vals' in d1
    assert type(d2) is dict and 'inds' in d2 and 'vals' in d2
    ###
    ### YOUR CODE HERE
    ###
    set1 = set(d1['inds'])
    set2 = set(d2['inds'])
    common_L = set1 & set2
    
    return sorted(list(common_L))

### Demo function call
d1 = {'inds': [9, 9, 1, 9, 8, 1], 'vals': [0.28, 0.84, 0.71, 0.03, 0.04, 0.75]}
d2 = {'inds': [0, 9, 9, 1, 3, 3, 9], 'vals': [0.26, 0.06, 0.46, 0.58, 0.42, 0.21, 0.53, 0.76]}
print(f'result={find_common_inds_test(d1,d2)}')
```

Output:
```
result=[1, 9]
```

 

**The demo should display this printed output.**
```
result=[1, 9]
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for find_common_inds_test (exercise 4). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_4, points: 1, grade: true, solution: false -->
```python
### Test Cell - Exercise 4  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=find_common_inds_test,
              ex_name='find_common_inds_test',
              key=b'5CpU4KJ6upxgu5f9QF_LK4KGaK5zkxGTtab-32LDWMo=', 
              n_iter=50)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to find_common_inds_test did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

Output:
```
initial memory usage: 0.00 MB
Test duration: 0.05 seconds
memory after test: 0.05 MB
memory peak during test: 1.25 MB
Passed! Please submit.
```

### FIN!
You've reached the end of this part. Don't forget to restart and run all cells again to make sure it's all working when run in sequence; and make sure your work passes the submission process. Good luck!