# Part 0: Representing numbers as strings

The following exercises are designed to reinforce your understanding of how we can view the encoding of a number as string of digits in a given base.

> If you are interested in exploring this topic in more depth, see the ["Floating-Point Arithmetic" section](https://docs.python.org/3/tutorial/floatingpoint.html) of the Python documentation.

```python
### Global imports
import dill
from cse6040_devkit import plugins, utils
from cse6040_devkit.training_wheels import run_with_timeout, suppress_stdout
import tracemalloc
from time import time


```

## Integers as strings

Consider the string of digits:

```python
    '16180339887'
```

If you are told this string is for a decimal number, meaning the base of its digits is ten (10), then its value is given by

$$
    [\![ \mathtt{16180339887} ]\!]_{10} = (1 \times 10^{10}) + (6 \times 10^9) + (1 \times 10^8) + \cdots + (8 \times 10^1) + (7 \times 10^0) = 16,\!180,\!339,\!887.
$$

Similarly, consider the following string of digits:

```python
    '100111010'
```

If you are told this string is for a binary number, meaning its base is two (2), then its value is

$$
    [\![ \mathtt{100111010} ]\!]_2 = (1 \times 2^8) + (1 \times 2^5) + \cdots + (1 \times 2^1).
$$

(What is this value?)

And in general, the value of a string of $d+1$ digits in base $b$ is,

$$
  [\![ s_d s_{d-1} \cdots s_1 s_0 ]\!]_b = \sum_{i=0}^{d} s_i \times b^i.
$$  

**Bases greater than ten (10).** Observe that when the base at most ten, the digits are the usual decimal digits, `0`, `1`, `2`, ..., `9`. What happens when the base is greater than ten? For this notebook, suppose we are interested in bases that are at most 36; then, we will adopt the convention of using lowercase Roman letters, `a`, `b`, `c`, ..., `z` for "digits" whose values correspond to 10, 11, 12, ..., 35.

### Exercise 0: (3 points)
**eval_string**  

**Your task:** define `eval_string` as follows:

Write a function, `eval_string(s, base)`. It takes a string of digits `s` in the base given by `base`. It returns its value as an integer.

That is, this function implements the mathematical object, $[\![ s ]\!]_b$, which would convert a string $s$ to its numerical value, assuming its digits are given in base $b$. For example:

```python
    eval_string('100111010', base=2) == 314
```

> Hint: Python makes this exercise very easy. Search Python's online documentation for information about the `int()` constructor to see how you can apply it to solve this problem. (You have encountered this constructor already, in Notebook/Assignment 2.)


```python
### Solution - Exercise 0  
def eval_string(s, base=2):
    assert type(s) is str
    assert 2 <= base <= 36
    return int(s, base)
    
    

### Demo function call
demo_ex0_tuples = [('100111010', 2), 
                   ('6040', 8),
                   ('deadbeef', 16),
                   ('4321', 5)]
results = []
for i, scenario in enumerate(demo_ex0_tuples):
    result = eval_string(scenario[0], scenario[1])
    print(f"eval_string({demo_ex0_tuples[i][0]}, {demo_ex0_tuples[i][1]})")
    print(f"--> {result}")
    results.append(result)


```

Output:
```
eval_string(100111010, 2)
--> 314
eval_string(6040, 8)
--> 3104
eval_string(deadbeef, 16)
--> 3735928559
eval_string(4321, 5)
--> 586
```

 

**The demo should display this printed output.**
```
eval_string(100111010, 2)
--> 314
eval_string(6040, 8)
--> 3104
eval_string(deadbeef, 16)
--> 3735928559
eval_string(4321, 5)
--> 586
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for eval_string (exercise 0). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_0, points: 3, grade: true, solution: false -->
```python
### Test Cell - Exercise 0  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=eval_string,
              ex_name='eval_string',
              key=b'ng8eQHSHSu3wLjeHp2gOHj42R3c7ii4n2k-Wqbz3OMU=', 
              n_iter=20)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to eval_string did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

Output:
```
initial memory usage: 0.00 MB
Test duration: 0.13 seconds
memory after test: 3.07 MB
memory peak during test: 4.07 MB
Passed! Please submit.
```

## Fractional values

Recall that we can extend the basic string representation to include a fractional part by interpreting digits to the right of the "fractional point" (i.e., "the dot") as having negative indices. For instance,

$$
    [\![ \mathtt{3.14} ]\!]_{10} = (3 \times 10^0) + (1 \times 10^{-1}) + (4 \times 10^{-2}).
$$

Or, in general,

$$
  [\![ s_d s_{d-1} \cdots s_1 s_0 \, \underset{\Large\uparrow}{\Huge\mathtt{.}} \, s_{-1} s_{-2} \cdots s_{-r} ]\!]_b = \sum_{i=-r}^{d} s_i \times b^i.
$$

### Exercise 1: (4 points)
**eval_strfrac**  

**Your task:** define `eval_strfrac` as follows:

Suppose a string of digits `s` in base `base` contains up to one fractional point. Complete the function, `eval_strfrac(s, base)`, so that it returns its corresponding floating-point value.

Your function should *always* return a value of type `float`, even if the input happens to correspond to an exact integer.

Examples:

```python
    eval_strfrac('3.14', base=10) ~= 3.14
    eval_strfrac('100.101', base=2) == 4.625
    eval_strfrac('2c', base=16) ~= 44.0   # Note: Must be a float even with an integer input!
```

> _Comment._ Because of potential floating-point roundoff errors, as explained in the videos, conversions based on the general polynomial formula given previously will not be exact. The testing code will include a built-in tolerance to account for such errors.
>
> _Hint._ You should be able to construct a solution that reuses the function, `eval_string()`, from Exercise 0.

```python
### Solution - Exercise 1  
def eval_strfrac(s, base=2):
    if '.' not in s:
        return float(eval_string(s,base))
    else:
        int_part,dec_part = s.split('.')
    int_eval = eval_string(int_part, base)
    # Convert fractional part
    val = 0
    for i, digit in enumerate(dec_part):
        val += int(digit, base) * (base ** -(i + 1))
    
    return float(int_eval + val)
        

### Demo function call
demo_ex1_tuples = [('3.14', 10), 
                   ('100.101', 2),
                   ('2c', 16),
                   ('2C', 16)]
results = []
for i, scenario in enumerate(demo_ex1_tuples):
    result = eval_strfrac(scenario[0], scenario[1])
    print(f"eval_strfrac({demo_ex1_tuples[i][0]}, {demo_ex1_tuples[i][1]})")
    print(f"--> {result}")
    results.append(result)


```

Output:
```
eval_strfrac(3.14, 10)
--> 3.14
eval_strfrac(100.101, 2)
--> 4.625
eval_strfrac(2c, 16)
--> 44.0
eval_strfrac(2C, 16)
--> 44.0
```

 

**The demo should display this printed output.**
```
eval_strfrac(3.14, 10)
--> 3.14
eval_strfrac(100.101, 2)
--> 4.625
eval_strfrac(2c, 16)
--> 44.0
eval_strfrac(2C, 16)
--> 44.0
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for eval_strfrac (exercise 1). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_1, points: 4, grade: true, solution: false -->
```python
### Test Cell - Exercise 1  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=eval_strfrac,
              ex_name='eval_strfrac',
              key=b'ng8eQHSHSu3wLjeHp2gOHj42R3c7ii4n2k-Wqbz3OMU=', 
              n_iter=20)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to eval_strfrac did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

Output:
```
initial memory usage: 0.00 MB
Test duration: 0.04 seconds
memory after test: 0.02 MB
memory peak during test: 1.25 MB
Passed! Please submit.
```

## Floating-point encodings

Recall that a floating-point encoding or format is a normalized scientific notation consisting of a _base_, a _sign_, a fractional _significand_ or _mantissa_, and a signed integer _exponent_. Conceptually, think of it as a tuple of the form, $(\pm, [\![s]\!]_b, x)$, where $b$ is the digit base (e.g., decimal, binary); $\pm$ is the sign bit; $s$ is the significand encoded as a base $b$ string; and $x$ is the exponent. For simplicity, let's assume that only the significand $s$ is encoded in base $b$ and treat $x$ as an integer value. Mathematically, the value of this tuple is $\pm \, [\![s]\!]_b \times b^x$.

**IEEE double-precision.** For instance, Python, R, and MATLAB, by default, store their floating-point values in a standard tuple representation known as _IEEE double-precision format_. It's a 64-bit binary encoding having the following components:

- The most significant bit indicates the sign of the value.
- The significand is a 53-bit string with an _implicit_ leading one. That is, if the bit string representation of $s$ is $s_0 . s_1 s_2 \cdots s_d$, then $s_0=1$ always and is never stored explicitly. That also means $d=52$.
- The exponent is an 11-bit string and is treated as a signed integer in the range $[-1022, 1023]$.

Thus, the smallest positive value in this format $2^{-1022} \approx 2.23 \times 10^{-308}$, and the smallest positive value greater than 1 is $1 + \epsilon$, where $\epsilon=2^{-52} \approx 2.22 \times 10^{-16}$ is known as _machine epsilon_ (in this case, for double-precision).

**Special values.** You might have noticed that the exponent is slightly asymmetric. Part of the reason is that the IEEE floating-point encoding can also represent several kinds of special values, such as infinities and an odd bird called "not-a-number" or `NaN`. This latter value, which you may have seen if you have used any standard statistical packages, can be used to encode certain kinds of floating-point exceptions that result when, for instance, you try to divide zero by zero.

> If you are familiar with languages like C, C++, or Java, then IEEE double-precision format is the same as the `double` primitive type. The other common format is single-precision, which is `float` in those same languages.

```python
def print_fp_hex(v):
    assert type(v) is float
    print("v = {} ==> v.hex() == '{}'".format(v, v.hex()))
    
print_fp_hex(0.0)
print_fp_hex(1.0)
print_fp_hex(16.0625)
print_fp_hex(-0.1)
```

Output:
```
v = 0.0 ==> v.hex() == '0x0.0p+0'
v = 1.0 ==> v.hex() == '0x1.0000000000000p+0'
v = 16.0625 ==> v.hex() == '0x1.0100000000000p+4'
v = -0.1 ==> v.hex() == '-0x1.999999999999ap-4'
```

**Inspecting a floating-point number in Python.** Python provides support for looking at floating-point values directly! Given any floating-point variable, `v` (that is, `type(v) is float`), the method `v.hex()` returns a string representation of its encoding. It's easiest to see by example, so run the following code cell:

Observe that the format has these properties:
* If `v` is negative, the first character of the string is `'-'`.
* The next two characters are always `'0x'`.
* Following that, the next characters up to but excluding the character `'p'` is a fractional string of hexadecimal (base-16) digits. In other words, this substring corresponds to the significand encoded in base-16.
* The `'p'` character separates the significand from the exponent. The exponent follows, as a signed integer (`'+'` or `'-'` prefix). Its implied base is two (2)---**not** base-16, even though the significand is.

Thus, to convert this string back into the floating-point value, you could do the following:
* Record the sign as a value, `v_sign`, which is either +1 or -1.
* Convert the significand into a fractional value, `v_signif`, assuming base-16 digits.
* Extract the exponent as a signed integer value, `v_exp`.
* Compute the final value as `v_sign * v_signif * (2.0**v_exp)`.

For example, here is how you can get 16.025 back from its `hex()` representation, `'0x1.0100000000000p+4'`:

```python
# Recall: v = 16.0625 ==> v.hex() == '0x1.0100000000000p+4'
print((+1.0) * eval_strfrac('1.0100000000000', base=16) * (2**4))
```

Output:
```
16.0625
```

### Exercise 2: (4 points)
**fp_bin**  

**Your task:** define `fp_bin` as follows:

Write a function, `fp_bin(v)`, that determines the IEEE-754 tuple representation of any double-precision floating-point value, `v`. That is, given the variable `v` such that `type(v) is float`, it should return a tuple with three components, `(s_sign, s_signif, v_exp)` such that

* `s_sign` is a string representing the sign bit, encoded as either a `'+'` or `'-'` character;
* `s_signif` is the significand, which should be a string of 54 bits having the form, `x.xxx...x`, where there are (at most) 53 `x` bits (0 or 1 values);
* `v_exp` is the value of the exponent and should be an _integer_.

For example:

```python
    v = -1280.03125
    assert v.hex() == '-0x1.4002000000000p+10'
    assert fp_bin(v) == ('-', '1.0100000000000010000000000000000000000000000000000000', 10)
```

> There are many ways to approach this problem. One we came up exploits the observation that $[\![\mathtt{0}]\!]_{16} == [\![\mathtt{0000}]\!]_2$ and $[\![\mathtt{f}]\!]_{16} = [\![\mathtt{1111}]\!]$ and applies an idea in this Stackoverflow post: https://stackoverflow.com/questions/1425493/convert-hex-to-binary  


```python
### Solution - Exercise 2  
def fp_bin(v):
    assert type(v) is float
    # Convert the float to hexadecimal '0x1.0100000000000p+4'
    hex_v = v.hex()
    
    # Extract the sign '+/-'
    if hex_v[0] != '-':
        v_sign = '+'
    else:
        v_sign = '-'
        
    # Extract significand, and exponent from the hexadecimal representation
 
    v_signif, v_exp = hex_v.split('p')
    
    v_signif = v_signif.replace('0x', '').replace('-', '').replace('+', '')
    
    # Convert significant to binary
    # 1. Split the hex into the leading bit and the fractional digits
    lead_bit, frac = v_signif.split('.')
   
    #v_signif = format(int(lead_bit,16), 'b') + '.' + (format(int(frac,16), '0' + str(len(frac)*4) + 'b') if frac else '0')
    
    v_signif = format(int(lead_bit, 16), 'b') + '.' + (format(int(frac, 16), '052b') if frac else '0' * 52)

    #convert an exponent to integer
    
    v_exp = int(v_exp)
    
    return v_sign, v_signif, v_exp
    
  
### Demo function call
demo_ex2_v = -1280.03125
result = fp_bin(demo_ex2_v)
print(result)

```

Output:
```
('-', '1.0100000000000010000000000000000000000000000000000000', 10)
```

 

**The demo should display this printed output.**
```
('-', '1.0100000000000010000000000000000000000000000000000000', 10)
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for fp_bin (exercise 2). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_2, points: 4, grade: true, solution: false -->
```python
### Test Cell - Exercise 2  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=fp_bin,
              ex_name='fp_bin',
              key=b'ng8eQHSHSu3wLjeHp2gOHj42R3c7ii4n2k-Wqbz3OMU=', 
              n_iter=21)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to fp_bin did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

Output:
```
initial memory usage: 0.00 MB
Test duration: 0.04 seconds
memory after test: 0.02 MB
memory peak during test: 1.24 MB
Passed! Please submit.
```

### Exercise 3: (2 points)
**eval_fp**  

**Your task:** define `eval_fp` as follows:

Suppose you are given a floating-point value in a base given by `base` and in the form of the tuple, `(sign, significand, exponent)`, where

* `sign` is either the character '+' if the value is positive and '-' otherwise;
* `significand` is a _string_ representation in base-`base`;
* `exponent` is an _integer_ representing the exponent value.

Complete the function,

```python
def eval_fp(sign, significand, exponent, base):
    ...
```

so that it converts the tuple into a numerical value (of type `float`) and returns it.

For example, `eval_fp('+', '1.25000', -1, base=10)` should return a value that is close to 0.125.


```python
### Solution - Exercise 3  
def eval_fp(sign, significand, exponent, base=2):
    assert sign in ['+', '-'], "Sign bit must be '+' or '-', not '{}'.".format(sign)
    assert type(exponent) is int
    if sign != '-':
        output = ((+1.0) * eval_strfrac(significand, base=base) * (base**exponent))
        #significand_val * (base**exponent)
    else:
        output = ((-1.0) * eval_strfrac(significand, base=base) * (base**exponent))
    return output

### Demo function call
demo_ex3_tuples = [('-', '2.G0EBPT', -1, 32), 
                   ('+', '1.20100202211020211211', 4, 3), 
                   ('+', '1.10101110101100100101001111000101', -5, 2)]
results = []
for i, scenario in enumerate(demo_ex3_tuples):
    result = eval_fp(scenario[0], scenario[1], scenario[2], scenario[3])
    print(f"eval_fp({demo_ex3_tuples[i][0]}, {demo_ex3_tuples[i][1]}, {demo_ex3_tuples[i][2]}, {demo_ex3_tuples[i][3]})")
    print(f"--> {result}")
    results.append(result)


```

Output:
```
eval_fp(-, 2.G0EBPT, -1, 32)
--> -0.0781387033930514
eval_fp(+, 1.20100202211020211211, 4, 3)
--> 138.25708894296503
eval_fp(+, 1.10101110101100100101001111000101, -5, 2)
--> 0.05257526742207119
```

 

**The demo should display this printed output.**
```
eval_fp(-, 2.G0EBPT, -1, 32)
--> -0.0781387033930514
eval_fp(+, 1.20100202211020211211, 4, 3)
--> 138.25708894296503
eval_fp(+, 1.10101110101100100101001111000101, -5, 2)
--> 0.05257526742207119
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for eval_fp (exercise 3). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_3, points: 2, grade: true, solution: false -->
```python
### Test Cell - Exercise 3  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=eval_fp,
              ex_name='eval_fp',
              key=b'ng8eQHSHSu3wLjeHp2gOHj42R3c7ii4n2k-Wqbz3OMU=', 
              n_iter=20)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to eval_fp did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

Output:
```
initial memory usage: 0.00 MB
Test duration: 0.04 seconds
memory after test: 0.03 MB
memory peak during test: 1.25 MB
Passed! Please submit.
```

### Exercise 4: (2 points)
**add_fp_bin**  

**Your task:** define `add_fp_bin` as follows:

Suppose you are given two binary floating-point values, `u` and `v`, in the tuple form given above. That is,
```python
    u == (u_sign, u_signif, u_exp)
    v == (v_sign, v_signif, v_exp)
```
where the base for both `u` and `v` is two (2). Complete the function `add_fp_bin(u, v, signif_bits)`, so that it returns the sum of these two values with the resulting significand _truncated_ to `signif_bits` digits.

For example:
```pythonA
u = ('+', '1.010010', 0)
v = ('-', '1.000000', -2)
assert add_fp_bin(u, v, 7) == ('+', '1.000010', 0)  # Caller asks for a significand with 7 digits
```
and:
```python
u = ('+', '1.00000', 0)
v = ('-', '1.00000', -6)
assert add_fp_bin(u, v, 6) == ('+', '1.11111', -1)  # Caller asks for a significand with 6 digits
```
(Check these examples by hand to make sure you understand the intended output.)

> _Note 0_: Assume that `signif_bits` _includes_ the leading 1. For instance, suppose `signif_bits == 4`. Then the significand will have the form, `1.xxx`.
>
> _Note 1_: You may assume that `u_signif` and `v_signif` use `signif_bits` bits (including the leading `1`). Furthermore, you may assume each uses far fewer bits than the underlying native floating-point type (`float`) does, so that you can use native floating-point to compute intermediate values.
>
> _Hint_: An earlier exercise defines a function, `fp_bin(v)`, which you can use to convert a Python native floating-point value (i.e., `type(v) is float`) into a binary tuple representation.


```python
### Solution - Exercise 4  
def add_fp_bin(u, v, signif_bits):
    u_sign, u_signif, u_exp = u
    v_sign, v_signif, v_exp = v
    # You may assume normalized inputs at the given precision, `signif_bits`.
    assert u_signif[:2] in {'1.', '0.'} and len(u_signif) == (signif_bits+1)
    assert v_signif[:2] in {'1.', '0.'} and len(v_signif) == (signif_bits+1)
    
    #convert the tuple into a numerical value (of type float)
    u_value = eval_fp(u_sign, u_signif, u_exp, base=2)
    v_value = eval_fp(v_sign, v_signif, v_exp, base=2)
    
    #Sum numerical values 
    sum_value = u_value + v_value
    
    #Convert the output back to tuple uisng fp_bin(v):
    
    sum_value_sign, sum_value_signif, sum_value_exp = fp_bin(sum_value)
    
    # Convert resulting significand truncated to signif_bits digits
    sum_value_signif = sum_value_signif[:signif_bits + 1]
    
    return (sum_value_sign, sum_value_signif, sum_value_exp)
    
    
### Demo function call
demo_ex4_tuples = [(('+', '1.010010', 0), ('-', '1.000000', -2), 7), 
                   (('+', '1.00000', 0), ('-', '1.00000', -6), 6)]
results = []
for i, scenario in enumerate(demo_ex4_tuples):
    result = add_fp_bin(scenario[0], scenario[1], scenario[2])
    print(f"add_fp_bin({demo_ex4_tuples[i][0]}, {demo_ex4_tuples[i][1]}, {demo_ex4_tuples[i][2]})")
    print(f"--> {result}")
    results.append(result)


```

Output:
```
add_fp_bin(('+', '1.010010', 0), ('-', '1.000000', -2), 7)
--> ('+', '1.000010', 0)
add_fp_bin(('+', '1.00000', 0), ('-', '1.00000', -6), 6)
--> ('+', '1.11111', -1)
```

 

**The demo should display this printed output.**
```
add_fp_bin(('+', '1.010010', 0), ('-', '1.000000', -2), 7)
--> ('+', '1.000010', 0)
add_fp_bin(('+', '1.00000', 0), ('-', '1.00000', -6), 6)
--> ('+', '1.11111', -1)
```


 ---
 <!-- Test Cell Boilerplate -->  
The cell below will test your solution for add_fp_bin (exercise 4). The testing variables will be available for debugging under the following names in a dictionary format.  
- `input_vars` - Input variables for your solution.   
- `original_input_vars` - Copy of input variables from prior to running your solution. Any `key:value` pair in `original_input_vars` should also exist in `input_vars` - otherwise the inputs were modified by your solution.  
- `returned_output_vars` - Outputs returned by your solution.  
- `true_output_vars` - The expected output. This _should_ "match" `returned_output_vars` based on the question requirements - otherwise, your solution is not returning the correct output. 


<!-- grade_id: ex_4, points: 2, grade: true, solution: false -->
```python
### Test Cell - Exercise 4  


from cse6040_devkit.tester_fw.testers import Tester
from yaml import safe_load
from time import time

tracemalloc.start()
mem_start, peak_start = tracemalloc.get_traced_memory()
print(f"initial memory usage: {mem_start/1024/1024:.2f} MB")

# Load testing utility
with open('resource/asnlib/publicdata/execute_tests', 'rb') as f:
    executor = dill.load(f)

@run_with_timeout(error_threshold=200.0, warning_threshold=100.0)
@suppress_stdout
def execute_tests(**kwargs):
    return executor(**kwargs)


# Execute test
start_time = time()
passed, test_case_vars, e = execute_tests(func=add_fp_bin,
              ex_name='add_fp_bin',
              key=b'ng8eQHSHSu3wLjeHp2gOHj42R3c7ii4n2k-Wqbz3OMU=', 
              n_iter=20)
# Assign test case vars for debugging
input_vars, original_input_vars, returned_output_vars, true_output_vars = test_case_vars
duration = time() - start_time
print(f"Test duration: {duration:.2f} seconds")
current_memory, peak_memory = tracemalloc.get_traced_memory()
print(f"memory after test: {current_memory/1024/1024:.2f} MB")
print(f"memory peak during test: {peak_memory/1024/1024:.2f} MB")
tracemalloc.stop()
if e: raise e
assert passed, 'The solution to add_fp_bin did not pass the test.'

###
### AUTOGRADER TEST - DO NOT REMOVE
###

print('Passed! Please submit.')
```

Output:
```
initial memory usage: 0.00 MB
Test duration: 0.04 seconds
memory after test: 0.02 MB
memory peak during test: 1.24 MB
Passed! Please submit.
```

**Fin.** If you have made it this far, congratulations on completing this part. **Don't forget to submit!**